<a href="https://colab.research.google.com/github/wkimunya/my/blob/main/Probability_Checkpoint_1_.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Probability and Combinatorics


```python
# Run this cell without changes
# You may find this module useful but are not required to use it
import math
```

## Probability

### Probability Exercises

Given the following probabilities:

$P(A) = 0.7$

$P(B) = 0.5$

$P(B|A) = 0.4$

Calculate the probabilities described in Question 1 and Question 2. Rather than "hard-coding" an answer, try to show your work by adding/subtracting/multiplying/dividing the original probabilities, e.g. `0.2*0.1` rather than just writing `0.02`

Hint: draw a diagram if you're getting stuck!

#### Question 1: $P(A and B)$
    
Assign your answer to the variable `ans1`

<!-- grade_id: cell-0124b2cebfd77e6c, points: 0, grade: false, solution: true -->
```python
#P(A and B) => P(A n B) => given b_a, we will multiply P(A)*P(B|A)
a = 0.7
b = 0.5
b_a = 0.4

ans1 = 0.7 * 0.4
# YOUR CODE HERE
ans1
```

Output:
```
0.27999999999999997
```

<!-- grade_id: bb63f014845b5b5cd3826c0fe8e8a4d5, points: 1, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# ans1 should be a floating point probability between 0 and 1
assert type(ans1) == float

```

#### Question 2: $P(A|B)$
    
Assign your answer to the variable `ans2`

<!-- grade_id: 35e11a8fcb18db65a5c0367f8b2cc322, points: 0, grade: false, solution: true -->
```python
# (A n B)/ B
# (A n B) =

ans2 = ans1/b
# YOUR CODE HERE

ans2
```

Output:
```
0.5599999999999999
```

<!-- grade_id: cell-279432d557f54c0a, points: 1, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# ans2 should be a floating point probability between 0 and 1
assert type(ans2) == float

```

### Probability Theory

#### Question 3: True or false: $P(A or B)$ is always equal to $P(A)$ + $P(B)$

<!-- grade_id: cell-3ac0d85c74130412, points: 0, grade: false, solution: true -->
```python
ans3 = True
# YOUR CODE HERE

ans3
```

Output:
```
True
```

<!-- grade_id: cell-0afd5e19ab559154, points: 1, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# ans3 should be a boolean
assert type(ans3) == bool

```

## Card Combinatorics

A standard deck of playing cards consists of 52 cards in each of the four suits of spades, hearts, diamonds, and clubs. Each suit contains 13 cards: Ace, 2, 3, 4, 5, 6, 7, 8, 9, 10, Jack, Queen, and King.
    
You have a standard deck of 52 cards and are asked the following questions:

Answer the questions below:

#### Question 4: What is the probability of drawing a King or a Queen?

Assign your answer to the variable `ans4`

<!-- grade_id: cell-7567c5aa3ef2dbcc, points: 0, grade: false, solution: true -->
```python
ans4 = 8/52
# YOUR CODE HERE

ans4
```

Output:
```
0.15384615384615385
```

<!-- grade_id: cell-06ca778aad6089a8, points: 1, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# ans4 should be a floating point probability between 0 and 1
assert type(ans4) == float

```

#### Question 5: How many possible 5-card combinations can be formed with this deck of 52 cards?

Assign your answer to the variable `ans5`

<!-- grade_id: 05176f852d6a7a9c5eca5d695467f412, points: 0, grade: false, solution: true -->
```python
#nCk = n!/(k!(n-k)!)

ans5 = (math.factorial(52)) /(math.factorial(5)*(math.factorial(52-5)))
# YOUR CODE HERE

ans5
```

Output:
```
2598960.0
```

<!-- grade_id: f5f6e2cb4ecc5832272d1ba4022c1db4, points: 1, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# ans5 should be a number greater than 1
assert type(ans5) == int or type(ans5) == float

```